# Lab 2: Inference in Graphical Models

### Machine Learning 2 (2017/2018)

* The lab exercises should be made in groups of two people.
* The deadline is Thursday, 29.04, 23:59.
* Assignment should be submitted through BlackBoard! Make sure to include your and your teammates' names with the submission.
* Attach the .IPYNB (IPython Notebook) file containing your code and answers. Naming of the file should be "studentid1\_studentid2\_lab#", for example, the attached file should be "12345\_12346\_lab1.ipynb". Only use underscores ("\_") to connect ids, otherwise the files cannot be parsed.

Notes on implementation:

* You should write your code and answers in an IPython Notebook: http://ipython.org/notebook.html. If you have problems, please ask.
* Use __one cell__ for code and markdown answers only!
    * Put all code in the cell with the ```# YOUR CODE HERE``` comment and overwrite the ```raise NotImplementedError()``` line.
    * For theoretical questions, put your solution using LaTeX style formatting in the YOUR ANSWER HERE cell.
* Among the first lines of your notebook should be "%pylab inline". This imports all required modules, and your plots will appear inline.
* Large parts of you notebook will be graded automatically. Therefore it is important that your notebook can be run completely without errors and within a reasonable time limit. To test your notebook before submission, select Kernel -> Restart \& Run All.

### Introduction
In this assignment, we will implement the sum-product and max-sum algorithms for factor graphs over discrete variables. The relevant theory is covered in chapter 8 of Bishop's PRML book, in particular section 8.4. Read this chapter carefuly before continuing!

We will implement sum-product and max-sum and apply it to a poly-tree structured medical diagnosis example.

For this assignment you should use numpy ndarrays (constructed with np.array, np.zeros, np.ones, etc.). We need n-dimensional arrays in order to store conditional distributions with more than one conditioning variable. If you want to perform matrix multiplication on arrays, use the np.dot function; all infix operators including *, +, -, work element-wise on arrays.

## Part 0: Doing the math (5 points)
We start with a set of three Bernoulli distributed variables X, Y, Z. Calculate the marginals $p(Y=1)$ and $p(Y=0)$ using only the sum and product rule where,

$$
p(X=1) = 0.05 \\\\
p(Z=1) = 0.2 \\\\
$$
$$
p(Y = 1 | X = 1, Z = 1) = 0.99 \\\\
p(Y = 1 | X = 1, Z = 0) = 0.9 \\\\
p(Y = 1 | X = 0, Z = 1) = 0.7 \\\\
p(Y = 1 | X = 0, Z = 0) = 0.0001 \\\\
$$

While implementing the message passing algorithms you should be able to use the results of this question as a guidance.

<!-- grade_id: cell-653065eb00c64d46, points: 5, grade: true, solution: true -->
We see when using sum rule, product rule, and that X,Z are independent:

\begin{align*}
P(Y=y)
&= \sum_{x=0}^1 \sum_{z=0}^1 p(Y=y,X=x,Z=z) \\
&= \sum_{x=0}^1 \sum_{z=0}^1 p(Y=y | X=x,Z=z) p(X=x,Z=z) \\
&= \sum_{x=0}^1 \sum_{z=0}^1 p(Y=y | X=x,Z=z) p(X=x)p(Z=z) \\
\end{align*}

This shows:
\begin{align*}
P(Y=1)
= 0.0001 \cdot (1-0.05) \cdot (1-0.2)
+ 0.7 \cdot (1-0.05) \cdot 0.2
+ 0.9 \cdot 0.05 \cdot (1-0.2)
+ 0.99 \cdot 0.05 \cdot 0.2
=0.178976
\end{align*}
from which it immediately follows that $p(Y=0) = 1- 0.178976 = 0.821024$.

## Part 1: The sum-product algorithm

We will implement a data structure to store a factor graph and to facilitate computations on this graph. Recall that a factor graph consists of two types of nodes, factors and variables. Below you will find some classes for these node types to get you started. Carefully inspect this code and make sure you understand what it does. Step by step will update its functionality.

```python
%pylab inline

class Node(object):
    """
    Base-class for Nodes in a factor graph. Only instantiate sub-classes of Node.
    """
    def __init__(self, name):
        # A name for this Node, for printing purposes
        self.name = name
        
        # Neighbours in the graph, identified with their index in this list.
        # i.e. self.neighbours contains neighbour 0 through len(self.neighbours) - 1.
        self.neighbours = []
        
        # Reset the node-state (not the graph topology)
        self.reset()
        
    def reset(self):
        # Incomming messages; a dictionary mapping neighbours to messages.
        # That is, it maps  Node -> np.ndarray.
        self.in_msgs = {}
        
        # A set of neighbours for which this node has pending messages.
        # We use a python set object so we don't have to worry about duplicates.
        self.pending = set([])

    def add_neighbour(self, nb):
        self.neighbours.append(nb)

    def send_sp_msg(self, other):
        # To be implemented in subclass.
        raise Exception('Method send_sp_msg not implemented in base-class Node')
   
    def send_ms_msg(self, other):
        # To be implemented in subclass.
        raise Exception('Method send_ms_msg not implemented in base-class Node')
    
    def receive_msg(self, other, msg):
        # Store the incomming message, replacing previous messages from the same node
        self.in_msgs[other] = msg

        # TODO: add pending messages
        # self.pending.update(...)
    
    def __str__(self):
        # This is printed when using 'print node_instance'
        return self.name


class Variable(Node):
    def __init__(self, name, num_states):
        """
        Variable node constructor.
        Args:
            name: a name string for this node. Used for printing. 
            num_states: the number of states this variable can take.
            Allowable states run from 0 through (num_states - 1).
            For example, for a binary variable num_states=2,
            and the allowable states are 0, 1.
        """
        self.num_states = num_states
        
        # Call the base-class constructor
        super(Variable, self).__init__(name)
    
    def set_observed(self, observed_state):
        """
        Set this variable to an observed state.
        Args:
            observed_state: an integer value in [0, self.num_states - 1].
        """
        # Observed state is represented as a 1-of-N variable
        # Set all-but-one states to a very low probability;
        # (a bit of a hack to avoid -inf values when taking logs)
        self.observed_state[:] = 1e-10
        self.observed_state[observed_state] = 1.0
        
    def set_latent(self):
        """
        Erase an observed state for this variable and consider it latent again.
        """
        # No state is preferred, so set all entries of observed_state to 1.0
        # Using this representation we need not differentiate observed an latent
        # variables when sending messages.
        self.observed_state[:] = 1.0
        
    def reset(self):
        super(Variable, self).reset()
        self.observed_state = np.ones(self.num_states)
        
    def marginal(self):
        """
        Compute the marginal distribution of this Variable.
        It is assumed that message passing has completed when this function is called.
        """
        # TODO: compute marginal
        return None
    
    def unnormalized_log_marginal(self):
        """
        Compute the unnormalized log marginal distribution of this Variable.
        It is assumed that message passing has completed when this function is called.
        """
        # TODO: compute unnormalized log marginal
        return None
    
    def send_sp_msg(self, other):
        # TODO: implement Variable -> Factor message for sum-product
        pass
   
    def send_ms_msg(self, other):
        # TODO: implement Variable -> Factor message for max-sum
        pass

class Factor(Node):
    def __init__(self, name, f, neighbours):
        """
        Factor node constructor.
        Args:
            name: a name string for this node. Used for printing
            f: a numpy.ndarray with N axes, where N is the number of neighbours.
               That is, the axes of f correspond to variables, and the index along that axes corresponds to a value of that variable.
               Each axis of the array should have as many entries as the corresponding neighbour variable has states.
            neighbours: a list of neighbouring Variables. Bi-directional connections are created.
        """
        # Call the base-class constructor
        super(Factor, self).__init__(name)
        
        f = np.array(f)  # For convenience, accept lists (of lists (of numbers or potentially lists of ...))
        
        assert len(neighbours) == f.ndim, 'Factor function f should accept as many arguments as this Factor node has neighbours'
        
        for nb_ind in range(len(neighbours)):
            nb = neighbours[nb_ind]
            assert f.shape[nb_ind] == nb.num_states, 'The range of the factor function f is invalid for input %i %s' % (nb_ind, nb.name)
            self.add_neighbour(nb)
            nb.add_neighbour(self)

        self.f = f
        
    def send_sp_msg(self, other):
        # TODO: implement Factor -> Variable message for sum-product
        pass
   
    def send_ms_msg(self, other):
        # TODO: implement Factor -> Variable message for max-sum
        pass

```

Output:
```
Populating the interactive namespace from numpy and matplotlib
```

### 1.1. Initialize the graph (5 points)
The equations in Part 0 can be represented by a factor graph. Instantiate this graph by creating Variable and Factor instances and linking them according to the graph structure. 
To instantiate the factor graph, first create the Variable nodes and then create Factor nodes, passing a list of neighbour Variables to each Factor. To get you started, we initialize the variable node for $X$ and the factor node corresponding to the prior $p(X)$.

<!-- grade_id: cell-85011c26139bed2e, points: 0, grade: false, solution: true -->
```python
X = Variable(name='X', num_states=2)
X_prior = Factor(name='p(X)',
                 f=np.array([0.95, 0.05]),
                 neighbours=[X])

# Please stick to the naming convention used below, otherwise the test functionality throughout the lab won't work
# Z =
# Z_prior =
                 
# Y =
# Y_cond =
                                 
# YOUR CODE HERE
#raise NotImplementedError()
Z = Variable(name='Z', num_states=2)
Z_prior = Factor(name='p(Z)', f=np.array([0.8,0.2]), neighbours=[Z])

Y = Variable(name='Y', num_states=2)
# Q[y][x][z] := p(Y=y|X=x,Z=z)
Q = [ [[0.9999,0.3],
       [0.1,0.01]],
     
      [[0.0001,0.7],
       [0.9,0.99]]
    ]
print(np.array(Q).shape)
Y_cond = Factor(name='p(Y|X,Z)', f=np.array(Q), neighbours=[Y,X,Z])

```

Output:
```
(2, 2, 2)
```

<!-- grade_id: cell-1bf4d184d405f181, points: 5, grade: true, solution: false -->
```python
### Test test test
assert Z_prior.f.shape == (2,)
assert Y_cond.f.shape == (2, 2, 2)


```

We will be doing a lot of marginalizations, i.e. obtain the distribution of a *single* variable from a distribution over multiple variables. Let's first write a function `marginalize` for that.

<!-- grade_id: cell-74ecd6f85f99a2d5, points: 0, grade: false, solution: true -->
```python
# Write a function marginalize that given
def marginalize(P, dim):
    # YOUR CODE HERE
    #raise NotImplementedError()
    marginal = np.zeros(P.shape[dim])
    for k in range(0, P.shape[dim]):
        marginal[k] = P.take(k, axis=dim).sum()
    return marginal
    
# Lets try it
test_P = np.random.rand(2, 3, 4)
test_P = test_P / test_P.sum()  # Normalize for proper distribution

# Do the marginal distributions look like you expect?
print (marginalize(test_P, 0))
print (marginalize(test_P, 1))
print (marginalize(test_P, 2))

```

Output:
```
[0.51860025 0.48139975]
[0.31212996 0.32304362 0.36482642]
[0.24325159 0.2790144  0.31431535 0.16341865]
```

### 1.2 Factor to variable messages (20 points)
Write a method `send_sp_msg(self, other)` for the Factor class, that checks if all the information required to pass a message to `Variable` `other` is present, computes the message and sends it to `other`. "Sending" here simply means calling the `receive_msg` function of the receiving node (we will implement this later). The message itself should be represented as a numpy array (np.array) whose length is equal to the number of states of the variable.

In the very end of 1.2 below we overwrite `send_sp_msg(self, other)` for the Factor class. In general, this is considered bad practise but in this lab it saves us from scrolling up and down all the time.

You will implement a function `send_sp_msg` that sends a message from a factor to a variable for the max-sum algorith. This function implements Equation 8.66 from Bishop. The message should be a numpy array (np.array) whose length is equal to the number of states of the variable.

It is a good idea to write a number of helper functions to implement these equations.

Since it is always a good idea to include checks, you should first write a method `can_send_message` that checks whether a node `node` has all the information required to pass a message to `other`. This should work for both variable and factor nodes.

<!-- grade_id: cell-4164822e95a6b68a, points: 0, grade: false, solution: true -->
```python
def can_send_message(sender, receiver):
    
    # YOUR CODE HERE
    #raise NotImplementedError()
    
    # sender can send msg to the receiver, if:
    # 1. sender is leaf node
    # OR if not leaf node:
    # 2. from all neighbours except receiver (n), sender has msg from n->sender
    
    # (extra) Case 0: receiver is not even a neighbour
    if receiver not in sender.neighbours:
        return False
    
    # Case 1: sender is leaf node
    if len(sender.neighbours) == 1:
        return True
    
    # Case 2: sender is not a leaf node
    for node in sender.neighbours:
        if node != receiver:
            # Check whether sender has got all msgs from that node
            if node not in sender.in_msgs.keys():
                return False
            
    return True
    

# Do the results make sense?
print (can_send_message(X, X_prior))
print (can_send_message(X_prior, X))

```

Output:
```
False
True
```


In Eq. 8.66, Bishop writes $f(x, x_1, ..., x_M)$, where $x$ corresponds to the variable that will receive the message. For now, assume this variable is the `index`-th neighbour of the factor. In order to ease implementation, it may be a good idea to write a function that rearanges the dimensions of `f` to comply with this notation, i.e. moves the dimension `index` to the front. Make sure to return a copy and keep all other dimensions in order! Use `np.moveaxis`.

<!-- grade_id: cell-702210e632c423bd, points: 0, grade: false, solution: true -->
```python
def move_dimension_first(f, index):
    f_new = f.copy()
    f_new = np.moveaxis(f_new, index, 0)
    return f_new
    
```

You should calculate the product of the incoming messages of all neighbours of the sending node except the receiving node. Therefore it may be useful to write a function `get_neighbour_messages` that gathers these messages in a list. If you want to not make things complicated, make sure the order of the messages in the list corresponds to the order of the variables in `neighbours`.

<!-- grade_id: cell-9e7f7986263e1543, points: 0, grade: false, solution: true -->
```python
def get_neighbour_messages(sender, receiver):
    list_of_msgs_vector = []
    
    for k in range(0,len(sender.neighbours)):
        node_k = sender.neighbours[k]
        if node_k != receiver:
            # (msg_list is vector/ndarray)
            msgs_vector = sender.in_msgs[node_k]
            list_of_msgs_vector.append(msgs_vector)
        
    return list_of_msgs_vector

```

Before marginalizing, we need to calculate $\prod_{m\in\text{ne}(f_s)\setminus x} \mu_{x_m\rightarrow f_s}(x_m)$ (Eq. 8.66) for all possible combinations of $x_1, ..., x_M$ (values of the neighbour nodes except the receiving node). An elegant and efficient way to calculate these is using the n-way outer product of vectors. This product takes n vectors $\mathbf{x}^{(1)}, \ldots, \mathbf{x}^{(n)}$ and computes a $n$-dimensional tensor (ndarray) whose element $i_0,i_1,...,i_n$ is given by $\prod_j \mathbf{x}^{(j)}_{i_j}$. In python, this is realized as `np.multiply.reduce(np.ix_(*vectors))` for a python list `vectors` of 1D numpy arrays. Try to figure out how this statement works -- it contains some useful functional programming techniques. What should `vectors` be? Try to see the link between the result and Eq. 8.66.

<!-- grade_id: cell-67269e06a554035c, points: 0, grade: false, solution: true -->
```python
def calc_other_neighbour_msg_prod(sender, receiver):
    
    # YOUR CODE HERE
    #raise NotImplementedError()
    list_of_msgs_vector = get_neighbour_messages(sender, receiver)    
    
    msgs_tensor = np.multiply.reduce(np.ix_(*list_of_msgs_vector))
    return msgs_tensor
    
```

Following Eq. 8.66, before marginalizing, you should calculate the product of $f(x, x_1, ..., x_M)$ with $\prod_{m\in\text{ne}(f_s)\setminus x} \mu_{x_m\rightarrow f_s}(x_m)$ for all configurations of $x, x_1, ..., x_M$. Since the second part does not depend on $x$, its tensor representations are of different dimensions. You can overcome this problem by using a loop, but preferably use numpy broadcasting by first aligning the dimensions of the tensors. You can use `np.expand_dims` or `X[None, ...]` to insert one dimension at the front. Write a function `calculate_factor` that, given `f` (which is reordered such that $x$ corresponds to the first dimension) and the (outer) product of the other neighbour messages, computes $f(x, x_1, ..., x_M) \prod_{m\in\text{ne}(f_s)\setminus x} \mu_{x_m\rightarrow f_s}(x_m)$ for all configurations of $x, x_1, ..., x_M$.

<!-- grade_id: cell-d34359928f3a997d, points: 0, grade: false, solution: true -->
```python
def calculate_factor(f_neighb_first, neighbour_msg_prod):
    
    # YOUR CODE HERE
    #raise NotImplementedError()
    expanded_neighbour_msg_prod = np.expand_dims(neighbour_msg_prod, axis=0)
    
    return f_neighb_first * expanded_neighbour_msg_prod
    
```

Put all the pieces together to define a function `calc_sum_product_factor_to_variable_msg` that calculates Eq. 8.66.

<!-- grade_id: cell-ec6b421860d40ae0, points: 0, grade: false, solution: true -->
```python
def calc_sum_product_factor_to_variable_msg(factor, variable):
    
    # YOUR CODE HERE
    #raise NotImplementedError()
    msgs_tensor = calc_other_neighbour_msg_prod(factor, variable)

    # calc index
    index_of_var = -1
    for k in range(0, len(factor.neighbours)):
        if factor.neighbours[k] == variable:
            index_of_var = k
            break

    #f_ordered is so that f_ordered = f(x,x_1,...,x_M) with x in first axis
    f_ordered = move_dimension_first(factor.f, index_of_var)
    
    f_mult_prod_msg_terms = calculate_factor(f_ordered, msgs_tensor)
    result = np.sum(f_mult_prod_msg_terms.reshape(variable.num_states,-1), axis=1)

    return result
```

```python
# Finally, we will define the send message function for you
def factor_send_sp_msg(self, variable):
    
    assert isinstance(variable, Variable), "Factor can only send messages to variable!"
    assert can_send_message(self, variable), "Cannot send message!"
    
    out_msg = calc_sum_product_factor_to_variable_msg(self, variable)
    
    # Send the message
    variable.receive_msg(self, out_msg)
    
    # Remove the pending sign if present
    self.pending.discard(variable)
    
Factor.send_sp_msg = factor_send_sp_msg

```

<!-- grade_id: cell-429fba08a94a3bd2, points: 20, grade: true, solution: false -->
```python
### Test test test
# message from X_prior to X
X_prior.reset()
X.reset()

X_prior.send_sp_msg(X)
assert np.allclose(list(X.in_msgs.values()), [0.95, 0.05])

# message from Z_prior to Z
Z_prior.reset()
Z.reset()

Z_prior.send_sp_msg(Z)
assert np.allclose(list(Z.in_msgs.values()), [0.8, 0.2])

# message from Y_cond to Y
Y_cond.reset()
Y.reset()

Y_cond.receive_msg(X, X_prior.f) # simulating that Y_cond received all necessary messages from X
Y_cond.receive_msg(Z, Z_prior.f) # simulating that Y_cond received all necessary messages from Z
Y_cond.send_sp_msg(Y)
assert np.allclose(list(Y.in_msgs.values()), [0.821024, 0.178976])

```

### 1.3 Variable to factor messages (10 points)

Write a method `calc_sum_product_variable_to_factor_msg(variable, factor)` that computes the message to be sent to a neighbour variable by a factor.

<!-- grade_id: cell-93d0a0d839eb99f3, points: 0, grade: false, solution: true -->
```python
def calc_sum_product_variable_to_factor_msg(variable, factor):
    
    # YOUR CODE HERE
    #raise NotImplementedError()
    
    out_msg = np.ones(variable.num_states)
    
    for k in range(0, variable.num_states):
        for neighbour_f in variable.neighbours:
            if neighbour_f != factor:
                out_msg[k] = out_msg[k] * variable.in_msgs[neighbour_f][k]
        
    return out_msg
    
```

```python
# Finally, we will define the send message function for you
def variable_send_sp_msg(self, factor):
    
    assert isinstance(factor, Factor), "Variable can only send messages to factor!"
    assert can_send_message(self, factor), "Cannot send message!"
    
    out_msg = calc_sum_product_variable_to_factor_msg(self, factor)
    
    # Send the message
    factor.receive_msg(self, out_msg)
    
    # Remove the pending sign if present
    self.pending.discard(factor)
    
Variable.send_sp_msg = variable_send_sp_msg

```

<!-- grade_id: cell-0cbe0bfa635450f3, points: 10, grade: true, solution: false -->
```python
### Test test test
Y_cond.reset()
Y.reset()

# First message from X to Y_cond
X_prior.reset()
X.reset()
X_prior.send_sp_msg(X) # simulating that X received all necessary messages
X.send_sp_msg(Y_cond) 
assert np.allclose(list(Y_cond.in_msgs.values()), [0.95, 0.05])

# Second message from Z to Y_cond
Z_prior.reset()
Z.reset()
Z_prior.send_sp_msg(Z) # simulating that Z received all necessary messages
Z.send_sp_msg(Y_cond)
assert np.allclose(list(Y_cond.in_msgs.values()), [[0.95, 0.05], [0.8, 0.2]])

```

### Testing a single forward pass and a single backward pass
Before we go on we will make sure that messages can be passed in both directions the graph.

```python
X_prior.reset()
X.reset()
Z_prior.reset()
Z.reset()
Y_cond.reset()
Y.reset()

# Forward pass
X_prior.send_sp_msg(X)
Z_prior.send_sp_msg(Z)
X.send_sp_msg(Y_cond)
Z.send_sp_msg(Y_cond)
Y_cond.send_sp_msg(Y)
assert np.allclose(list(Y.in_msgs.values()), [0.821024, 0.178976])

# Backward pass
Y.send_sp_msg(Y_cond)
Y_cond.send_sp_msg(X)
Y_cond.send_sp_msg(Z)
X.send_sp_msg(X_prior)
Z.send_sp_msg(Z_prior)
assert np.allclose(list(X.in_msgs.values()), [[0.95, 0.05],[1., 1.]])
assert np.allclose(list(Z.in_msgs.values()), [[0.8, 0.2],[1., 1.]])

```

### 1.4 Compute marginal (10 points)
Later in this assignment, we will implement message passing schemes to do inference. Once the message passing has completed, we will want to compute local marginals for each variable.
Write the method `marginal` for the `Variable` class, that computes a marginal distribution over that node.

<!-- grade_id: cell-0f6434f5c5d4dc9c, points: 0, grade: false, solution: true -->
```python
def marginal(self):
    
    # YOUR CODE HERE
    #raise NotImplementedError()
    marginal = np.ones(self.num_states)
    for k in range(0, self.num_states):
        for neighbour_f in self.neighbours:
            marginal[k] = marginal[k] * self.in_msgs[neighbour_f][k]
    
    # in cases where factor functions werent normalized
    marginal = marginal / sum(marginal)
    return marginal

Variable.marginal = marginal

```

```python
# NOTE: We added this extra block, because for marginals of every node to be computed with the algorithm,
# a forward and a backward pass is needed. Therefore, in this check we do the backward pass too.
# Otherwise only the marginal of Y could have been checked.
#
# The cell block after this (the original one) only has the forward pass towards variable Y,
# but checks for the marginal p(X).

X_prior.reset()
X.reset()
Z_prior.reset()
Z.reset()
Y_cond.reset()
Y.reset()

# Forward pass
X_prior.send_sp_msg(X)
Z_prior.send_sp_msg(Z)
X.send_sp_msg(Y_cond)
Z.send_sp_msg(Y_cond)
Y_cond.send_sp_msg(Y)

# Backward pass
Y.send_sp_msg(Y_cond)
Y_cond.send_sp_msg(X)
Y_cond.send_sp_msg(Z)
X.send_sp_msg(X_prior)
Z.send_sp_msg(Z_prior)

print("X.marginal() is: " + str(X.marginal()))
print("Z.marginal() is: " + str(Z.marginal()))
print("Y.marginal() is: " + str(Y.marginal()))
assert np.allclose(X.marginal(), [0.95, 0.05])
assert np.allclose(Z.marginal(), [0.8, 0.2])
assert np.allclose(Y.marginal(), [0.821024, 0.178976])

```

Output:
```
X.marginal() is: [0.95 0.05]
Z.marginal() is: [0.8 0.2]
Y.marginal() is: [0.821024 0.178976]
```

<!-- grade_id: cell-7f666a68bbbe81e5, points: 10, grade: true, solution: false -->
```python
### Test test test
# Simulate a single forward pass
X_prior.reset()
X.reset()
Z_prior.reset()
Z.reset()
Y_cond.reset()
Y.reset()

X_prior.send_sp_msg(X)
Z_prior.send_sp_msg(Z)
X.send_sp_msg(Y_cond)
Z.send_sp_msg(Y_cond)
Y_cond.send_sp_msg(Y)

assert np.allclose(X.marginal(), [0.95, 0.05])
assert np.allclose(Z.marginal(), [0.8, 0.2])
assert np.allclose(Y.marginal(), [0.821024, 0.178976])

```

Output:
```
<ipython-input-19-9fd61ec70e4c> in <module>()
     14 Y_cond.send_sp_msg(Y)
     15 
---> 16 assert np.allclose(X.marginal(), [0.95, 0.05])
     17 assert np.allclose(Z.marginal(), [0.8, 0.2])
     18 assert np.allclose(Y.marginal(), [0.821024, 0.178976])

<ipython-input-17-32b16392bf95> in marginal(self)
      6     for k in range(0, self.num_states):
      7         for neighbour_f in self.neighbours:
----> 8             marginal[k] = marginal[k] * self.in_msgs[neighbour_f][k]
      9 
     10     # in cases where factor functions werent normalized

KeyError: <__main__.Factor object at 0x7ff5740cf7b8>
```

### 1.5 Receiving messages (10 points)
In order to implement a message passing algorithms, we need some way to determine which nodes are ready to send messages to which neighbours. We make use of the concept of "pending messages", which is explained in Bishop (8.4.7): 
"we will say that a (variable or factor)
node a has a message pending on its link to a node b if node a has received any
message on any of its other links since the last time it sent a message to b. Thus,
when a node receives a message on one of its links, this creates pending messages
on all of its other links."

Before we say node a has a pending message for node b, we **must check that node a has received all messages needed to compute the message that is to be sent to b**.

Modify the function `receive_msg`, so that it updates the self.pending variable as described above. The member self.pending is a set that is to be filled with Nodes to which self has pending messages.

<!-- grade_id: cell-39b7b3a84cdf1e56, points: 0, grade: false, solution: true -->
```python
# ANSWER 1.5
def receive_msg(self, other, msg):
    self.in_msgs[other] = msg
    
    for nodei in self.neighbours: #check for all neighbouring nodes different from 
        if nodei != other:
            if can_send_message(self, nodei): #check if node has received all messages
                self.pending.update([nodei]) #add pending node
    
Node.receive_msg = receive_msg

```

<!-- grade_id: cell-42c321c98bf9b6b1, points: 10, grade: true, solution: false -->
```python
### Test test test
X_prior.reset()
X.reset()
assert X_prior.pending == set()

X_prior.pending.add(X)
assert str(list(X_prior.pending)[0]) == X.name

X_prior.send_sp_msg(X)
assert X_prior.pending == set()

```

### 1.6 Inference Engine (10 points)
Write a function `sum_product(node_list)` that runs the sum-product message passing algorithm on a tree-structured factor graph with given nodes. The input parameter `node_list` is a list of all Node instances in the graph, which is assumed to be ordered correctly. That is, the list starts with a leaf node, which can always send a message. Subsequent nodes in `node_list` should be capable of sending a message when the pending messages of preceding nodes in the list have been sent. The sum-product algorithm then proceeds by passing over the list from beginning to end, sending all pending messages at the nodes it encounters. Then, in reverse order, the algorithm traverses the list again and again sends all pending messages at each node as it is encountered. For this to work, we initialized pending messages for all the leaf nodes, e.g. `X_prior.pending.add(X)`, where `X_prior` is a Factor node corresponding the the prior, `X` is a `Variable` node and the only connection of `X_prior` goes to `X`.

<!-- grade_id: cell-6a6fb2568d840c68, points: 0, grade: false, solution: true -->
```python
def sum_product(node_list):
#    for node in node_list:
#        for pending in node.pending.copy():
#            node.send_sp_msg(pending)


    for node in node_list:
        for target_node in node.pending.copy():
            node.send_sp_msg(target_node)

    for node in reversed(node_list):
        for target_node in node.pending.copy():
            node.send_sp_msg(target_node)

```

<!-- grade_id: cell-fb8315046f8fde42, points: 10, grade: true, solution: false -->
```python
### Test test test
nodes = [X_prior, X, Z_prior, Z, Y_cond, Y]
for n in nodes:
    n.reset()
    
X_prior.pending.add(X)
Z_prior.pending.add(Z)
Y.pending.add(Y_cond)

sum_product(nodes)
assert np.allclose(Y.marginal(), [0.821024, 0.178976])

```

### 1.7 Observed variables and probabilistic queries (15 points)
We will now use the inference engine to answer probabilistic queries. That is, we will set certain variables to observed values, and obtain the marginals over latent variables. We have already provided functions `set_observed` and `set_latent` that manage a member of Variable called `observed_state`. Modify the `calc_sum_product_variable_to_factor_msg` and `Variable.marginal` routines that you wrote before, to use `observed_state` so as to get the required marginals when some nodes are observed.

<!-- grade_id: cell-55c50fb9c1836789, points: 0, grade: false, solution: true -->
```python
def calc_sum_product_variable_to_factor_msg(variable, factor):
    
    # YOUR CODE HERE
#     raise NotImplementedError()
    
    out_msg = np.ones(variable.num_states)
    
    for k in range(0, variable.num_states):
        for neighbour_f in variable.neighbours:
            if neighbour_f != factor:
                out_msg[k] = out_msg[k] * variable.in_msgs[neighbour_f][k] * variable.observed_state[k]
        
    return out_msg
    
```

<!-- grade_id: cell-13801e018b6f9d27, points: 10, grade: true, solution: false -->
```python
### Test, test, test
X_prior.reset()
X.reset()
Y_cond.reset()

X_prior.send_sp_msg(X)
X.set_observed(0)
X.send_sp_msg(Y_cond)
assert np.allclose(list(Y_cond.in_msgs.values()), [9.5e-01, 5.0e-12])

```

<!-- grade_id: cell-d0b91da808874e73, points: 0, grade: false, solution: true -->
```python
def marginal(self):
    
    # YOUR CODE HERE
#     raise NotImplementedError()


    marginal = self.observed_state
    for key, value in self.in_msgs.items():
        
        marginal = marginal * value
    
    marginal = marginal / np.sum(marginal)
        
    return marginal

   # marginal = np.ones(self.num_states)
   # for k in range(0, self.num_states):
   #     for neighbour_f in self.neighbours:
#        marginal[k] = marginal[k] * self.in_msgs[neighbour_f][k] * self.observed_state[k]
    
    # in cases where factor functions werent normalized
 #   marginal = marginal / sum(marginal)
  #  return marginal

Variable.marginal = marginal

```

<!-- grade_id: cell-2f8e5d9052ed73ef, points: 5, grade: true, solution: false -->
```python
### Test, test, test
# Simulate a single forward pass
X_prior.reset()
X.reset()
Z_prior.reset()
Z.reset()
Y_cond.reset()
Y.reset()

X.set_observed(0)
Z.set_observed(0)

X_prior.send_sp_msg(X)
Z_prior.send_sp_msg(Z)
X.send_sp_msg(Y_cond)
Z.send_sp_msg(Y_cond)
Y_cond.send_sp_msg(Y)

assert np.allclose(Y.marginal(), [9.99900000e-01, 1.00000022e-04])

```

## Part 2: The max-sum algorithm
Next, we implement the max-sum algorithm as described in section 8.4.5 of Bishop.

### 2.1 Factor to variable messages (10 points)
Implement the function `Factor.send_ms_msg` that sends Factor -> Variable messages for the max-sum algorithm. It is analogous to the `Factor.send_sp_msg` function you implemented before. Make sure it works for observed and unobserved nodes. Consider using a number of helper functions as seen in Part 1.

<!-- grade_id: cell-389621edefd825a3, points: 0, grade: false, solution: true -->
```python
def calc_other_neighbour_msg_sum(sender, receiver):
    list_of_msgs_vector = get_neighbour_messages(sender, receiver)    
    
    msgs_tensor = np.sum(np.ix_(*list_of_msgs_vector))
    return msgs_tensor

def calculate_factor_ms(f_neighb_first, neighbour_msg_sum):
    expanded_neighbour_msg_sum = np.expand_dims(neighbour_msg_sum, axis=0)
    
    return np.log(f_neighb_first) + expanded_neighbour_msg_sum

                                    
def calc_max_sum_factor_to_variable_msg(factor, variable):
    msgs_tensor = calc_other_neighbour_msg_sum(factor, variable)

    # calc index
    index_of_var = -1
    for k in range(0, len(factor.neighbours)):
        if factor.neighbours[k] == variable:
            index_of_var = k
            break

    #f_ordered is so that f_ordered = f(x,x_1,...,x_M) with x in first axis
    f_ordered = move_dimension_first(factor.f, index_of_var)
    
    f_mult_prod_msg_terms = calculate_factor_ms(f_ordered, msgs_tensor)
    result = np.max(f_mult_prod_msg_terms.reshape(variable.num_states,-1), axis=1)

    return result
                                                 

def factor_send_ms_msg(self, variable):
     
  
    out_msg = calc_max_sum_factor_to_variable_msg(self, variable)
        
    # Send the message
    variable.receive_msg(self, out_msg)
    
    # Remove the pending sign if present
    self.pending.discard(variable)
    
Factor.send_ms_msg = factor_send_ms_msg

```

<!-- grade_id: cell-c6156e1505633bc4, points: 10, grade: true, solution: false -->
```python
### Test test test
# message from X_prior to X
X_prior.reset()
X.reset()

X_prior.send_ms_msg(X)
assert np.allclose(list(X.in_msgs.values()), [-0.05129329, -2.99573227])

# message from Z_prior to Z
Z_prior.reset()
Z.reset()

Z_prior.send_ms_msg(Z)
assert np.allclose(list(Z.in_msgs.values()), [-0.22314355, -1.60943791])

# message from Y_cond to Y
Y_cond.reset()
Y.reset()

Y_cond.receive_msg(X, X_prior.f) # simulating that Y_cond received all necessary messages from X
Y_cond.receive_msg(Z, Z_prior.f) # simulating that Y_cond received all necessary messages from Z
Y_cond.send_ms_msg(Y)
assert np.allclose(list(Y.in_msgs.values()), [1.74989999, 0.79332506])

```

### 2.2 Variable to factor messages (10 points)
Implement the `Variable.send_ms_msg` function that sends Variable -> Factor messages for the max-sum algorithm.

<!-- grade_id: cell-183948f06650a7e0, points: 0, grade: false, solution: true -->
```python
# ANSWER 2.2
def calc_max_sum_variable_to_factor_msg(variable, factor):
    
    # YOUR CODE HERE
    # raise NotImplementedError()

    out_msg = np.zeros(variable.num_states)
    
    for k in range(0, variable.num_states):
        for neighbour_f in variable.neighbours:
            if neighbour_f != factor:
                out_msg[k] = out_msg[k] + variable.in_msgs[neighbour_f][k]+np.log(variable.observed_state[k])
        
    return out_msg
    
def variable_send_ms_msg(self, factor):
    assert isinstance(factor, Factor), "Variable can only send messages to factor!"
    assert can_send_message(self, factor), "Cannot send message!"
    out_msg = calc_max_sum_variable_to_factor_msg(self, factor)

    # YOUR CODE HERE
    
  #  raise NotImplementedError()
    
    # Send the message
    factor.receive_msg(self, out_msg)
    
    # Remove the pending sign if present
    self.pending.discard(factor)

Variable.send_ms_msg = variable_send_ms_msg

```

<!-- grade_id: cell-bbf08d58e9b309c9, points: 10, grade: true, solution: false -->
```python
### Test test test
Y_cond.reset()
Y.reset()

# First message from X to Y_cond
X_prior.reset()
X.reset()
X_prior.send_ms_msg(X) # simulating that X received all necessary messages
X.send_ms_msg(Y_cond)
assert np.allclose(list(Y_cond.in_msgs.values()), [-0.05129329, -2.99573227])

# Second message from Z to Y_cond
Z_prior.reset()
Z.reset()
Z_prior.send_ms_msg(Z) # simulating that Z received all necessary messages
Z.send_ms_msg(Y_cond)
assert np.allclose(list(Y_cond.in_msgs.values()), [[-0.05129329, -2.99573227], [-0.22314355, -1.60943791]])

```

### 2.3 Implement unnormalized log marginal (5 points)
Write the method `unnormalized_log_marginal` for the `Variable` class, that computes a unnormalized log marginal distribution over that node.

<!-- grade_id: cell-96de1de3d9b1da94, points: 0, grade: false, solution: true -->
```python
def unnormalized_log_marginal(self):
    
    # YOUR CODE HERE    
    marginal = np.log(self.observed_state)
    for key, value in self.in_msgs.items():                
        marginal += value    
    
    unnormalized_log_marginal = marginal
    
    return unnormalized_log_marginal

Variable.unnormalized_log_marginal = unnormalized_log_marginal
```

<!-- grade_id: cell-d0c08fd4b0857dde, points: 5, grade: true, solution: false -->
```python
### Test test test
# Simulate a single forward pass
X_prior.reset()
X.reset()
Z_prior.reset()
Z.reset()
Y_cond.reset()
Y.reset()

X_prior.send_ms_msg(X)
Z_prior.send_ms_msg(Z)
X.send_ms_msg(Y_cond)
Z.send_ms_msg(Y_cond)
Y_cond.send_ms_msg(Y)

assert np.allclose(X.unnormalized_log_marginal(), [-0.05129329, -2.99573227])
assert np.allclose(Z.unnormalized_log_marginal(), [-0.22314355, -1.60943791])
assert np.allclose(Y.unnormalized_log_marginal(), [-0.27453685, -2.01740615])

```

### 2.4 Find a MAP state (10 points)

Using the same message passing schedule we used for sum-product, implement the max-sum algorithm. For simplicity, we will ignore issues relating to non-unique maxima. So there is no need to implement backtracking; the MAP state is obtained by a per-node maximization (eq. 8.98 in Bishop). Make sure your algorithm works with both latent and observed variables.

<!-- grade_id: cell-f94559550902ac16, points: 0, grade: false, solution: true -->
```python
def max_sum(node_list):
    for node in node_list:
        for pending in node.pending.copy():
            node.send_ms_msg(pending)
    for node in reversed(node_list):
        for target_node in node.pending.copy():
            node.send_sp_msg(target_node)

    # YOUR CODE HERE
    #raise NotImplementedError()

```

<!-- grade_id: cell-ac2d26dcb27a8d75, points: 5, grade: true, solution: false -->
```python
### Test test test: unobserved
nodes = [X_prior, X, Z_prior, Z, Y_cond, Y]
for n in nodes:
    n.reset()
    
X_prior.pending.add(X)
Z_prior.pending.add(Z)
Y.pending.add(Y_cond)

max_sum(nodes)

assert np.allclose(Y.unnormalized_log_marginal(), [-0.27453685, -2.01740615] )

```

<!-- grade_id: cell-f06b8b4bec85b4dd, points: 5, grade: true, solution: false -->
```python
### Test test test: partiallY observed
nodes = [X_prior, X, Z_prior, Z, Y_cond, Y]
for n in nodes:
    n.reset()
    
X_prior.pending.add(X)
Z_prior.pending.add(Z)
Y.pending.add(Y_cond)

Z.set_observed(1)

max_sum(nodes)

assert np.allclose(Y.unnormalized_log_marginal(), [-2.86470401, -2.01740615])

```

Given the max-marginals what do you have to do in order to find the global optimum? Why can we neglect the normalization constant here?

<!-- grade_id: cell-605e729ee910dbc8, points: 5, grade: true, solution: true -->
For each marginal, we need to find the set of $x_i^*$ values which maximizes the marginal.
According to the equation under equation (8.90) in Bishop, we need to multiply all previous sets of marginals and maximise this result. WIth such large amount of product some problems such as over flow might inccur. Therefore, we need to take the logarithm of the results. This operation makes the original max-product equations become the sum-max equation. The normalization constant $Z$ becomes zero and we neglect it from the calculation.

### Part 3: Medical graph
Now that we implemented the sum-product and max-sum algorithm. We will apply them to a poly-tree structured medical diagnosis example.


### 3.1 Initialize the graph (5 points)

Convert the directed graphical model ("Bayesian Network") shown below to a factor graph. Instantiate this graph by creating Variable and Factor instances and linking them according to the graph structure. 
To instantiate the factor graph, first create the Variable nodes and then create Factor nodes, passing a list of neighbour Variables to each Factor.
Use the following prior and conditional probabilities.

$$
p(\verb+Influenza+) = 0.05 \\\\
p(\verb+Smokes+) = 0.2 \\\\
$$
$$
p(\verb+SoreThroat+ = 1 | \verb+Influenza+ = 1) = 0.3 \\\\
p(\verb+SoreThroat+ = 1 | \verb+Influenza+ = 0) = 0.001 \\\\
p(\verb+Fever+ = 1| \verb+Influenza+ = 1) = 0.9 \\\\
p(\verb+Fever+ = 1| \verb+Influenza+ = 0) 0.05 \\\\
p(\verb+Bronchitis+ = 1 | \verb+Influenza+ = 1, \verb+Smokes+ = 1) = 0.99 \\\\
p(\verb+Bronchitis+ = 1 | \verb+Influenza+ = 1, \verb+Smokes+ = 0) = 0.9 \\\\
p(\verb+Bronchitis+ = 1 | \verb+Influenza+ = 0, \verb+Smokes+ = 1) = 0.7 \\\\
p(\verb+Bronchitis+ = 1 | \verb+Influenza+ = 0, \verb+Smokes+ = 0) = 0.0001 \\\\
p(\verb+Coughing+ = 1| \verb+Bronchitis+ = 1) = 0.8 \\\\
p(\verb+Coughing+ = 1| \verb+Bronchitis+ = 0) = 0.07 \\\\
p(\verb+Wheezing+ = 1| \verb+Bronchitis+ = 1) = 0.6 \\\\
p(\verb+Wheezing+ = 1| \verb+Bronchitis+ = 0) = 0.001 \\\\
$$

<!-- grade_id: cell-216d36a73771fdb7, points: 5, grade: true, solution: true -->
```python
=Variable(name='', num_states=2)
_prior=Factor(name='p()',
                 f=np.array([0.95, 0.05]),
                 neighbours=[])


Smokes = Variable(name='Smokes', num_states=2)
Smokes_prior = Factor(name='p(Smokes)', f=np.array([0.8,0.2]), neighbours=[Smokes])

Q = [ [[0.9999,0.3],
       [0.1,0.01]],
     
      [[0.0001,0.7],
       [0.9,0.99]]
    ]

 = Variable(name='', num_states=2)
_cond = Factor(name='p(|,Smokes)', f=np.array(Q), neighbours=[, Influenza, Smokes])


=Variable(name='', num_states=2)
_cond = Factor(name='p(|Influenza)', f=np.array([[0.999,0.7],[0.001,0.3]]), neighbours=[, ])

=Variable(name='', num_states=2)
_cond = Factor(name='p(|Influenza)', f=np.array([[0.95,0.1],[0.05,0.9]]), neighbours=[, ])

=Variable(name='', num_states=2)
_cond = Factor(name='p(|)', f=np.array([[0.93, 0.2],[0.07,0.8]]), neighbours=[, ])

Wheezing=Variable(name='Wheezing', num_states=2)
Wheezing_cond = Factor(name='p(Wheezing|)', f=np.array([[0.999, 0.4],[0.001, 0.6]]), neighbours=[Wheezing, ])

```

### 3.2. Sum-product algorithm (10 points)

3.2.1 Run the sum-product algorithm on an unobserved graph. Print the marginal for every variable node.

<!-- grade_id: cell-07681153d48af573, points: 5, grade: true, solution: true -->
```python

# YOUR CODE HERE
#raise NotImplementedError()

# An ordering (where Bronchitis is root node):
nodes = [_prior, Smokes_prior, Smokes, SoreThroat, SoreThroat_cond, Fever, Fever_cond, Influenza, Bronchitis_cond, Wheezing, Wheezing_cond, , Coughing_cond, Bronchitis]
for n in nodes:
    n.reset()

# Pending the linked nodes
_prior.pending.add()
Smokes_prior.pending.add(Smokes)
.pending.add(_cond)
.pending.add(_cond)
.pending.add(_cond)
Wheezing.pending.add(Wheezing_cond)
# Calculate the sum product
sum_product(nodes)

# Printing the output nodes
print("I.marginal() is: " + str(I.marginal()))
print("Smokes.marginal() is: " + str(Smokes.marginal()))
print(".marginal() is: " + str(.marginal()))
print(".marginal() is: " + str(.marginal()))
print("Fever.marginal() is: " + str(Fever.marginal()))
print(".marginal() is: " + str(.marginal()))
print("Wheezing.marginal() is: " + str(Wheezing.marginal()))

```

Output:
```
I.marginal() is: [0.95 0.05]
Smokes.marginal() is: [0.8 0.2]
.marginal() is: [0.821024 0.178976]
.marginal() is: [0.98405 0.01595]
Fever.marginal() is: [0.9075 0.0925]
.marginal() is: [0.79934752 0.20065248]
Wheezing.marginal() is: [0.89179338 0.10820662]
```

3.2.2 Rerun the sum-product algorithm on an partially observed graph, where the variable 'Influenza' is set to 0. Print the marginal for every variable node.

<!-- grade_id: cell-83214d3cc28cb7db, points: 5, grade: true, solution: true -->
```python
for n in nodes:
    n.reset()

# Pending the linked nodes
_prior.pending.add()
Smokes_prior.pending.add(Smokes)
.pending.add(_cond)
.pending.add(_cond)
.pending.add(_cond)
Wheezing.pending.add(Wheezing_cond)
#Set the 'Influenza' variable to 0
Influenza.set_observed(0)
# Calculate the sum product
sum_product(nodes)
    
# Printing the output nodes
print("I.marginal() is: " + str(I.marginal()))
print("Smokes.marginal() is: " + str(Smokes.marginal()))
print(".marginal() is: " + str(.marginal()))
print(".marginal() is: " + str(.marginal()))
print("Fever.marginal() is: " + str(Fever.marginal()))
print(".marginal() is: " + str(.marginal()))
print("Wheezing.marginal() is: " + str(Wheezing.marginal()))
    
# YOUR CODE HERE
#raise NotImplementedError()

```

Output:
```
I.marginal() is: [1.00000000e+00 5.26315789e-12]
Smokes.marginal() is: [0.8 0.2]
.marginal() is: [0.85992 0.14008]
.marginal() is: [0.999 0.001]
Fever.marginal() is: [0.95 0.05]
.marginal() is: [0.8277416 0.1722584]
Wheezing.marginal() is: [0.91509208 0.08490792]
```

### 3.3. max-sum algorithm (10 points)

3.3.1 Run the max_sum algorithm on an unobserved graph. Print the marginal for every variable node.

<!-- grade_id: cell-e5088e48afe006b9, points: 5, grade: true, solution: true -->
```python
for n in nodes:
    n.reset()

# Pending the linked nodes
_prior.pending.add()
Smokes_prior.pending.add(Smokes)
.pending.add(_cond)
.pending.add(_cond)
.pending.add(_cond)
Wheezing.pending.add(Wheezing_cond)

# Calculate the max sum
max_sum(nodes)
    
# Printing the output nodes
print("I.unnormalized_log_marginal() is: " + str(I.unnormalized_log_marginal()))
print("Smokes.unnormalized_log_marginal() is: " + str(Smokes.unnormalized_log_marginal()))
print(".unnormalized_log_marginal() is: " + str(.unnormalized_log_marginal()))
print(".unnormalized_log_marginal() is: " + str(.unnormalized_log_marginal()))
print("Fever.unnormalized_log_marginal() is: " + str(Fever.unnormalized_log_marginal()))
print(".unnormalized_log_marginal() is: " + str(.unnormalized_log_marginal()))
print("Wheezing.unnormalized_log_marginal() is: " + str(Wheezing.unnormalized_log_marginal()))
    

#raise NotImplementedError()

```

Output:
```
I.unnormalized_log_marginal() is: [ 0.77526154 -2.13807448]
Smokes.unnormalized_log_marginal() is: [2.09402896 0.96112911]
.unnormalized_log_marginal() is: [-0.40040184 -2.80366912]
.unnormalized_log_marginal() is: [0.29388604 0.1249635 ]
Fever.unnormalized_log_marginal() is: [0.14105242 1.26908841]
.unnormalized_log_marginal() is: [-0.82098808 -2.08736864]
Wheezing.unnormalized_log_marginal() is: [-1.31613934 -1.3761055 ]
```

3.3.2 Rerun the max_sum algorithm on an partially observed graph, where the variable 'Influenza' is set to 0. Print the marginal for every variable node.

<!-- grade_id: cell-f74e50a37f4e31b9, points: 5, grade: true, solution: true -->
```python
for n in nodes:
    n.reset()

# Pending the linked nodes
_prior.pending.add()
Smokes_prior.pending.add(Smokes)
.pending.add(_cond)
.pending.add(_cond)
.pending.add(_cond)
Wheezing.pending.add(Wheezing_cond)
#Set the 'Influenza' variable to 0
Influenza.set_observed(0)
# Calculate the maxsum
max_sum(nodes)
    
# Printing the output nodes
print("I.unnormalized_log_marginal() is: " + str(I.unnormalized_log_marginal()))
print("Smokes.unnormalized_log_marginal() is: " + str(Smokes.unnormalized_log_marginal()))
print(".unnormalized_log_marginal() is: " + str(.unnormalized_log_marginal()))
print(".unnormalized_log_marginal() is: " + str(.unnormalized_log_marginal()))
print("Fever.unnormalized_log_marginal() is: " + str(Fever.unnormalized_log_marginal()))
print(".unnormalized_log_marginal() is: " + str(.unnormalized_log_marginal()))
print("Wheezing.unnormalized_log_marginal() is: " + str(Wheezing.unnormalized_log_marginal()))
    

#raise NotImplementedError()

```

Output:
```
I.unnormalized_log_marginal() is: [  0.77526154 -25.16392541]
Smokes.unnormalized_log_marginal() is: [48.23295575 51.20573678]
.unnormalized_log_marginal() is: [-0.40040184 -2.80366912]
.unnormalized_log_marginal() is: [2.30994028e-03 2.31225254e-06]
Fever.unnormalized_log_marginal() is: [4.28465161e-05 2.25507979e-06]
.unnormalized_log_marginal() is: [-0.82098808 -2.08736864]
Wheezing.unnormalized_log_marginal() is: [-1.31613934 -1.3761055 ]
```

```python

```

```python

```

```python

```

```python

```